---
title: "ENSC311 Lab10: More penguins"
tutorial:
  id: ENSC311_Lec10
output: learnr::tutorial
runtime: shiny_prerendered
---
<!-- File created by tutorialize.  Do not edit by hand. -->


```{r setup, include=FALSE}
library(learnr)
knitr::opts_chunk$set(echo = FALSE)
gradethis::gradethis_setup()
palmer_p <- palmerpenguins::penguins

### I added the lines below
options(tutorial.event_recorder = tutorialize::recorder)
options(tutorial.storage = learnr::filesystem_storage(file.path(path.expand("~"), "_learnr_storage_")))

```



```{r filter_mean-setup}
#```{r palmer-p-setup}
# If you name this "filter_mean-setup" it will automatically be the setup for 
# filter-mean chunk
palmer_p <- palmerpenguins::penguins
```


Lets see what islands have the most large penguins. In the sand box below:

- Create a data frame named large_penguins.
- Use the `select()` function to select body mass, species, island, and sex from the palmer_p data frame.
- Use the `%>%` operator and the `drop_na()` function to remove Nas from the data set
- Use the `filter` function to select penguins with an above average body mass.

```{r filter_mean, exercise=TRUE}

```

```{r filter_mean-solution}
# ```{r filter_mean, exercise = TRUE, exercise.setup = palmer-p}
large_penguins <- palmer_p %>%
  select(species, island, body_mass_g, sex)%>%
  drop_na()%>%
  filter(body_mass_g > mean(body_mass_g))
```

```{r filter_mean-code-check}
grade_code()
```
Use the data frame you. just created and the sand box below to:

- Make a box plot with island on the x axis and body mass on the y axis.
 
```{r filter_mean_plot-setup}
#```{r palmer-p-setup}
# If you name this "filter_mean-setup" it will automatically be the setup for 
# filter-mean chunk
palmer_p <- palmerpenguins::penguins
large_penguins <- palmer_p %>% select(species, island, body_mass_g, 
    sex) %>% drop_na() %>% filter(body_mass_g > mean(body_mass_g))
```

```{r filter_mean_plot, exercise=TRUE}

```

```{r filter_mean_plot-solution}
#```{r filter_mean_plot, exercise = true, exercise.setup = filter_mean}
ggplot(large_penguins, aes(island, body_mass_g))+
  geom_boxplot()
```

```{r filter_mean_plot-code-check}
grade_code()
```

This plot is helpful but there may be something skewing the data.  use the sand box below to:

- Color the plot based on the penguins species

```{r box_with_color, exercise=TRUE}

```

```{r box_with_color-solution}
#```{r box_with_color, exercise = TRUE}
ggplot(large_penguins, aes(island, body_mass_g, color = species))+
  geom_boxplot()
```

```{r box_with_color-code-check}
grade_code()
```

Lets look at how bill length changes with species.  In the sand box below:

- create a box plot with species on the x axis and bill length on the y axis
- use the. `facet_wrap()` function to separate the graph by penguin sex
- finally color the graph based of species of penguin

```{r bill_length, exercise=TRUE}

```

```{r bill_length-solution}
# ```{r bill_length, exercise = TRUE}
ggplot(palmer_p, aes(species,bill_length_mm, color = species))+
  geom_boxplot()+
  facet_wrap(facets = "sex")
```

```{r bill_length-code-check}
grade_code()
```

lets make a new data frame.  In the sand box below:

- use the `select()` function to select flipper length, body mass, species, sex, and island from the plamer_p data frame
- Name your data frame b_f

```{r final_plot_df-setup}
#```{r palmer-p-setup}
# If you name this "filter_mean-setup" it will automatically be the setup for 
# filter-mean chunk
palmer_p <- palmerpenguins::penguins
large_penguins <- palmer_p %>% select(species, island, body_mass_g, 
    sex) %>% drop_na() %>% filter(body_mass_g > mean(body_mass_g))
```

```{r final_plot_df, exercise=TRUE}

```

```{r final_plot_df-solution}
#```{r final_plot_df, exercise = TRUE}
b_f <- palmer_p %>%
  select(flipper_length_mm, body_mass_g, species, island,sex)
```

```{r final_plot_df-code-check}
grade_code()
```

Use your data frame in the sand box below:

- to create a scatter plot of body mass and flipper length
- facet and color by sex

```{r facet_by_sex, exercise=TRUE}

```

```{r facet_by_sex-solution}
# ```{r facet_by_sex, exercise = TRUE}
ggplot(b_f, aes(body_mass_g, flipper_length_mm, color = sex))+
  geom_point()+
  facet_wrap(facets = "sex")
```

```{r facet_by_sex-code-check}
grade_code()
```

this plot is okay but lets break it down more. In the sand box below:

- to create a scatter plot of body mass and flipper length
- facet and color by species make sure to set scales = to free 
- use the `geom_smooth()` function to better visualize the data.
- use. the `labs` function to change the legend title to Species, the main title to Flipper Length Vs. Body Mass, the x axis to Body Mass (g) and the y axis to Flipper Length (mm).

```{r final_plot, exercise=TRUE}

```

```{r final_plot-solution}
# ```{r final_plot, exercise = TRUE}
ggplot(b_f, aes(body_mass_g, flipper_length_mm, color = species))+
  geom_point()+
  geom_smooth()+
  facet_wrap(facets = "species", scales = "free")+
  labs(color = "Species", title = "Flipper Length Vs. Body Mass", x = "Body Mass (g)", y = "Flipper Length (mm)")
```

```{r final_plot-code-check}
grade_code()
```
