---
title: "ENSC311 Lec 08: Working with lists and data.frames"
tutorial:
  id: "ENSC311_Lec08"
output: 
  learnr::tutorial:
    progressive: true
runtime: shiny_prerendered
description: "An introduction to data.frames -- indexing (subsetting) tabular data in base R."

---
<!-- File created by tutorialize.  Do not edit by hand. -->

```{r setup, include=FALSE}
knitr::opts_chunk$set(echo = FALSE)
library(learnr)
gradethis::gradethis_setup()
options(tutorial.event_recorder = tutorialize::mongo_recorder)
options(tutorial.storage = learnr::filesystem_storage(file.path(path.expand("~"), "_learnr_storage_")))

options(tutorial.events_collection = 
  mongolite::mongo(
    db = "tutorial",
    collection = "events",
    url = persistR::getPersistent("db_key", "~/R/persistR")))

```

## R Objects: A review

```{r, echo=TRUE}
dir()
```


In the class so far, we've come to understand that everthing in R is an `object`.  We've worked with three types of objects so far -- vectors, functions, and object names.  When we assign an object (like a vector or function) to an object name, we can recall that object using the name:

```{r, echo=TRUE}
example_name <- 1:10
example_name
```

or

```{r, echo=TRUE}
area_of_circle <- function(radius) {
  # pi is a built-in name with the value of 3.1415927...
  area <- pi * radius^2
  return(area)
}

area_of_circle(6.3)
```

We've also worked with `atomic vectors`, which are ordered collections of values and we've come to understand that there is a sub-type of `atomic vectors` for each different types of data.  So far we've learned about `numeric`, `Date`, `character`, and `logical` data types.  We've also found that `numeric` atomic vectors are one of two flavors: `integer` or `double`.  (A `double` contains real numbers -- those with a decimal point like 3.256)





We refer to these relationships as "inheritance."  We say "`numeric` inherits from `atomic`".  By this be mean that `numerics` have all of the properties of `atomic`s, plus a few more.  More importantly, if you can do a particular operation on an object type, you can do that same operation on any object that inherits from that type.  

For instance, the `+` operator is designed to work on `numeric` vectors.  Therefore, it will work with either `integers` or `doubles`.  Importantly, the `[` operator (for indexing) is designed to work on vectors, so it works on any type of vector.

## The other type of vector: a `list`.

So, if an `atomic` vector is an ordered collection of values, what is a `list` vector? Well, as `list` vector (which is usually just called a `list`) is an ordered collection of `objects`.  

A useful analogy for a `list` is a line of cubbyholes.  The line of cubbyholes is itself a single object, but you can put any type object in each cubbyhole! 

Try making a list below.  We'll use the `list` function to create a list of length 3 (analogous to a line of three cubbyholes).  

See if you can fill in the blanks so that the list has:

- the `letters` vector in the first position,
- a vector of `integers` ranging from 1 to 10 in the second position, and 
- the `sum` function in the third position.

Don't forget to display the list!

```{r}
# make my_list for use in examples
my_list <- list(letters, 1:10, sum)
```

```{r 01_make_a_list, exercise=TRUE}
my_list <- list(___, ___, ___)
```

```{r 01_make_a_list-solution}
my_list <- list(letters, 1:10, sum)
my_list
```

```{r 01_make_a_list-code-check}
grade_code()
```

The last item in the list might look a little strange.  But in the sandbox below: 

- display the `sum` function and see if it doesn't look like the last item in `my_list`!

```{r 02_display_sum, exercise=TRUE}

```

```{r 02_display_sum-solution}
sum
```

```{r 02_display_sum-code-check}
grade_code()
```

### Indexing lists by position

Remember indexing?  Indexing is using the `[` operator to return a subset of a vector, like so:

```{r, echo = TRUE}
best_university <- LETTERS[c(13, 19, 21)] 
best_university
```

Index with the `[` operator works on any vector.  Therefore, indexing works on `atomic` vectors or `list`s, because both are `vectors`!  

There are one important thing to remember, however, when indexing a `list`.

When applied to a list, the `[` operator will *ALWAYS* return a list.

That point bears repeating...

When applied to a list, the `[` operator will *ALWAYS* return a list.

And one last time...

When applied to a list, the `[` operator will *ALWAYS* return a list.

Even though I said it three times, you may forget.  At least most people do.  It's important.  Try to remember.

So earlier in the tutorial we created `my_list`.

- Remind yourself what is in `my_list` by displaying it

```{r 03_display_my_list-setup}
my_list <- list(letters, 1:10, sum)
```

```{r 03_display_my_list, exercise=TRUE}

```

```{r 03_display_my_list-solution}
my_list
```

```{r 03_display_my_list-code-check}
grade_code()
```

OK.  Now, using your knowlege of indexing `atomic` vectors:

- index `my_list` by position to create a new list that has only `letters` and the `sum` function; assign the new list to `my_short_list`
- display `my_short_list`
- prove to yourself that `my_short_list` is a list (hint: try `?is.list` in the console).

```{r 04_shorten_my_list-setup}
my_list <- list(letters, 1:10, sum)
```

```{r 04_shorten_my_list, exercise=TRUE}

```

```{r 04_shorten_my_list-solution}
my_short_list <- my_list[c(1,3)]
my_short_list
is.list(my_short_list)
```

```{r 04_shorten_my_list-code-check}
grade_code()
```

### A list of length 1

Try the following:

- index `my_list` by position to create a new list that has only the integer vector from `my_list`; assign the new list to `length_1_list`
- display `length_1_list`
- prove to yourself that `length_1_list` is a list.

```{r 05_length_1_list-setup}
my_list <- list(letters, 1:10, sum)
```

```{r 05_length_1_list, exercise=TRUE}

```

```{r 05_length_1_list-solution}
length_1_list <- my_list[2]
length_1_list
is.list(length_1_list)
```

```{r 05_length_1_list-code-check}
grade_code()
```

In the sandbox below:

- Try to sum the values in `length_1_list` by passing `length_1_list` to the `sum` function (you can't submit the answer for a reason...)

```{r 06_cant_sum_a_list-setup}
my_list <- list(letters, 1:10, sum)
length_1_list <- my_list[2]
```

```{r 06_cant_sum_a_list, exercise=TRUE}

```

So what gives?  Why doesn't `sum` work?

```{r, 07_sum_a_list_quiz}
quiz(
  question(
    "Why can't you sum the values in `length_1_list` by passing it to the `sum` function?",
    answer("Because the numeric vector is still inside a `list`.  You are trying to sum a list. You can't sum a `list`!", correct = TRUE)
  ))
```

Hopefully, you got the right answer on that last one... ;-)

## Extracting from lists

So what good is indexing something from a list if you can't `sum` it?  For instance, how do you work with the `atomic` vector that is stuck inside `my_list`?

You can extract an item from a list using one of the following:

1) the `[[` operator (extract by position or by name)
2) the `$` operator (extract by name)

Let's look at how these work

### The `[[` operator

Extracting an item from a list looks a lot like indexing, but it is fundamentally different.  

If we go back to our cubbyhole analogy, `[` retrieves a collection of cubbyholes (along with the contents of the cubbyholes) from a line of cubbyholes.  In the example below, we use indexing to create a collection containing only one cubbyhole -- in other words, a list of length = 1.

```{r, echo=TRUE}
length_1_list <- my_list[2]
is.list(length_1_list)
length(length_1_list)
```

Extraction, on the other hand, retrieves the *contents* of one cubbyhole!  In the sandbox below, I've given you an example of how to extract the integer vector from `my_list` and assigning it to `my_integer_vector`.

- Prove to yourself that `my_integer_vector` is not a list
- Prove to yourself that `my_integer_vector` is a vector
- Have R display the length of `my_integer_vector`.
- Sum the values in `my_integer_vector`

```{r 08_extract_integer_vector-setup}
my_list <- list(letters, 1:10, sum)
```

```{r 08_extract_integer_vector, exercise=TRUE}
# Note the double brackets...
my_integer_vector <- my_list[[2]]
```

```{r 08_extract_integer_vector-solution}
# Note the double brackets...
my_integer_vector <- my_list[[2]]
is.list(my_integer_vector)
is.vector(my_integer_vector)
length(my_integer_vector)
sum(my_integer_vector)
```

```{r 08_extract_integer_vector-code-check}
grade_code()
```

The `identical()` function tells you if two objects are identical.  Consider the following:

```{r, echo=TRUE}
my_list[2]
1:10
identical(1:10, my_list[2])
```

Why is `my_list[2]` not identical to `1:10`?  Because `my_list[2]` yields a list of length = 1 containing the vector `1:10` (in other words, a vector that is still inside a cubbyhole).  Note carefully the difference in how R displays the result of `my_list[2]` and `1:10`.  The "[[1]]" gives you a hint that R is displaying a `list`.

Remember, the extraction operator `[[` will pull the `integer` vector out of the cubbyhole (extract it from the `list`).  In the sandbox below, run the sample code to see what it returns. Then:

- modify `my_list[2]` in the sample line of code slightly so that `identical` returns `TRUE`

```{r 09_try_extraction-setup}
my_list <- list(letters, 1:10, sum)
```

```{r 09_try_extraction, exercise=TRUE}
identical(1:10, my_list[2])
```

```{r 09_try_extraction-solution}
identical(1:10, my_list[[2]])
```

```{r 09_try_extraction-code-check}
grade_code()
```

### Naming the objects in a `list`

Just like each value in an `atomic` vector can be assigned a name, the items in a list can be assigned names.  In the sandbox below...

- Use your knowledge of naming `atomic` vectors to see if you can deduce how to name the objects in `my_list` as "letters", "integers", and "sum_function".
- Then, display `my_list`

```{r}
# add names to my_list for code examples
names(my_list) <- c("letters", "integers", "sum_function")
```

```{r 10_name_my_list-setup}
my_list <- list(letters, 1:10, sum)
```

```{r 10_name_my_list, exercise=TRUE}

```

```{r 10_name_my_list-solution}
names(my_list) <- c("letters", "integers", "sum_function")
my_list
```

```{r 10_name_my_list-code-check}
grade_code()
```

Notice the "[[#]]" preceding each list item is gone; now the name is displayed preceded by a `$` when R displays each item in a named `list`.

Also, note that we can name objects in a list when we create the `list`, simply by putting including the names in the `list()` function.  The syntax looks like this:

```{r}
a_named_list <- 
  list(
    person = c("Carol", "Shelly", "Sam"),
    age = c(57, 35, 23),
    height = c(5.4, 5.2, 5.9),
    is_a_student = c(FALSE, TRUE, TRUE))
a_named_list
```

In the sandbox below:

- change the code so that the names "letters", "integers", and "sum_function" are included in `my_list` at the time the `list` is created.
- display `my_list`

```{r 11_name_when_creating_list, exercise=TRUE}
my_list <- list(letters, 1:10, sum)
```

```{r 11_name_when_creating_list-solution}
my_list <- list(letters = letters, integers = 1:10, sum_function = sum)
my_list
```

```{r 11_name_when_creating_list-code-check}
grade_code()
```

In the sandbox below:

- see if you can deduce how to *index* the letters and the `sum` function from `my_list` using names rather than position.  The result should be a list of length = 2. (Don't save the resulting `list`, just let R display it).
- then *index* the `integer` vector from `my_list`, again using a name rather than a position (the result should be a `list` of length = 1 containing `1:10`)
- then see if you can *extract* the `integer` vector from `my_list` using names rather than position.  The result should be an `integer` vector of length = 10.

```{r 12_index_extract_by_name-setup}
my_list <- list(letters, 1:10, sum)
names(my_list) <- c("letters", "integers", "sum_function")
```

```{r 12_index_extract_by_name, exercise=TRUE}

```

```{r 12_index_extract_by_name-solution}
my_list[c("letters", "sum_function")]
my_list["integers"]
my_list[["integers"]]
```

```{r 12_index_extract_by_name-code-check}
grade_code()
```

### Using the `$` operator

Extracting an item from a list is a very common thing to do.  And admittedly, typing `my_list[["integers"]]` is a bit cumbersome for such a common task.  So R provides a bit of a shortcut: the `$` operator.  The `$` operator only works on named lists.  You put the name of the list in front of the `$` and the name of the object you want after the `$`.  Unexpectedly, the name of the object is *not* put in quotes.

In other words, `my_list$integers` will extract the "integers" vector from `my_list`.  Consider:

```{r}
my_list$integers
is.list(my_list$integers)
is.vector(my_list$intergers)
length(my_list$integers)
sum(my_list$integers)
```

In the sandbox below:

- extract the "integers" vector from `my_list` using first the `$` operator and then again using the `[[` operator; show the results are identical

```{r 13_dollar_vs_double_square-setup}
my_list <- list(letters, 1:10, sum)
names(my_list) <- c("letters", "integers", "sum_function")
```

```{r 13_dollar_vs_double_square, exercise=TRUE}

```

```{r 13_dollar_vs_double_square-solution}
identical(my_list$integers, my_list[["integers"]])
```

```{r 13_dollar_vs_double_square-code-check}
grade_code()
```

Now consider the following code:

```{r, echo=TRUE}
bob <- my_list$sum_function
bob(my_list[["integers"]])
```

```{r 14_extract_function_quiz}
quiz(
  question(
    "Why did the code above result in the value `55`?",
    answer('Because the sum of the "integers" vector in `my_list` is 55.'),
    answer('Because the "sum_function" member of `my_list` contains a copy of the `sum` function.'),
    answer('Because we assigned a copy of the `sum` function to the name `bob` and then passed the "integers" vector from list to the `bob` function.'),
    answer('All of the above.', correct = TRUE)
  ))
```

Finally, see if the following code makes sense.  Remember, `(` instructs R to call the function found before the `(` by passing the arguments found between the `(` and the `)`.

```{r, echo=TRUE}
my_list$sum_function(my_list$integers)
```

and

```{r, echo=TRUE}
my_list[[3]](my_list[[2]])
```

Pretty wild, huh?

### R likes to makes copies.

With only one exception (which we will explore later in the course), R always make copies of `objects` when they are assigned to a name (using `<-`) or placed in a `list`.

So realize that we don't actually have the `letters` vector in `my_list`.  We have a *copy* of the `letters` vector.  Similarly, we don't have the `sum` function, but a *copy* of the sum function.  This means that we can play with or modify objects contained in a `list` without fear of modifying the original objects from which they may have been derived.

## Data.frames

So we have reached the point where you have all of the knowledge you need to start working with real data sets.  In R, data sets are stored in objects called `data.frames`.

`data.frames` contain tabular data.  In other words, `data.frames` have data arranged in rows and columns.

R has some built-in `data.frames`.  Consider `mtcars`:

```{r}
mtcars
```

Each row in `mtcars` represents a make and model of car and each column represents a characteristic of the cars.  Can you imagine how you might learn more about the `mtcars` data set from the console?  (Yes, you guessed it... enter `?mtcars` in the console!)

Do that now and read about the data.

### Exploring data.frames

The `str()` function provides a useful summary of the structure of a data.frame.  In the sandbox below:

- pass `mtcars` to the `str()` function

```{r 15_str_mtcars, exercise=TRUE}

```

```{r 15_str_mtcars-solution}
str(mtcars)
```

```{r 15_str_mtcars-code-check}
grade_code()
```

The `head()` function shows the first few values of any item.  In the sandbox below:

- pass `letters` to the `head()` function;
- pass `mtcars` to the `head()` function.

```{r 16_head_letters_mtcars, exercise=TRUE}

```

```{r 16_head_letters_mtcars-solution}
head(letters)
head(mtcars)
```

```{r 16_head_letters_mtcars-code-check}
grade_code()
```

Read the help on the head function.  Now, in the sandbox below:

- use `head()` to display the first 10 letters of the alphabet;
- use `head()` to display the forst 10 rows of `mtcars`.

```{r 17_head_letters_mtcars_10, exercise=TRUE}

```

```{r 17_head_letters_mtcars_10-solution}
head(letters, 10)
head(mtcars, 10)
```

```{r 17_head_letters_mtcars_10-code-check}
grade_code()
```

`nrow()` and `ncol()` are two useful functions `data.frames`.  Check out the help on these function.  Then, in the sandbox below:

- display the number of rows in `mtcars`
- display the number columns

```{r 18_nrow_ncol, exercise=TRUE}

```

```{r 18_nrow_ncol-solution}
nrow(mtcars)
ncol(mtcars)
```

```{r 18_nrow_ncol-code-check}
grade_code()
```

## `data.frames` are `list`s

What's important to realize, here, is that `data.frames` are `lists`.  They `lists` with superpowers, but they are `lists` none-the-less.  

In other words, `data.frames` can do more than a regular `list`.  However, since a `data.frame` is a `list`, everything you just learned about `lists` applies to a `data.frame`.

The secret?  Just remember that a `data.frame` is a `list` of `vectors`; each `vector` represents a column in the `data.frame`.

I'll say that again.  

A `data.frame` is a `list` of `vectors`; each `vector` represents a column in the `data.frame`.

Keep reading the prior sentence until it really sinks in...

In the sandbox below:

- Prove that `mtcars` is a `list`

```{r 19_df_is_list, exercise=TRUE}

```

```{r 19_df_is_list-solution}
is.list(mtcars)
```

```{r 19_df_is_list-code-check}
grade_code()
```

So, if a `data.frame` is a `list` of `vectors` that each represent columns, then what would the names of the `vectors` in the `list` represent? Do the following and then compare the results from each command in the sandbox:

- Use `names()` to get the names of the items in `mtcars`
- Then use `str()` to view the structure of `mtcars` again.

```{r 20_names_of_mtcars, exercise=TRUE}

```

```{r 20_names_of_mtcars-solution}
names(mtcars)
str(mtcars)
```

```{r 20_names_of_mtcars-code-check}
grade_code()
```

Sure! The names of the items in the `list` must be the column names for the `data.frame`!  And so they are!

How about the length of the `list`?

- Display the length of `mtcars` using `length()`
- Display the number of columns useing `ncol()`

```{r 21_length_vs_ncol, exercise=TRUE}

```

```{r 21_length_vs_ncol-solution}
length(mtcars)
ncol(mtcars)
```

```{r 21_length_vs_ncol-code-check}
grade_code()
```

The number of items in (length of) the `list` is equal to the number of columns in the `data.frame`.  Makes sense, right?

### Extracting columns from a data.frame

If `mtcars` is a list of columns, then you should be able to extract columns.  Try it!

- use the `$` operator to extract the `mpg` column from `mtcars`; assign the result to `mpg_col`.
- display `mpg_col`
- show that `mpg_col` is a vector
- use `[[` to extract the 6th column from `mtcars`; assign the result to `col_6`.
- display `col_6`
- show that `col_6` is a vector

```{r 22_extract_columns, exercise=TRUE}

```

```{r 22_extract_columns-solution}
mpg_col <- mtcars$mpg
mpg_col
is.vector(mpg_col)
col_6 <- mtcars[[6]]
col_6
is.vector(col_6)
```

```{r 22_extract_columns-code-check}
grade_code()
```

Remember, `mpg_col` contains a *copy* of the data in the `mpg` column from `mtcars`.  If you modify `mpg_col`, it doesn't affect `mtcars`.

How long are the column vectors?

- display the length of the `mpg_col` vector
- display the length of the `col_6` vector
- display the number of rows in `mtcars`

```{r 23_col_length_vs_nrows-setup}
mpg_col <- mtcars$mpg
col_6 <- mtcars[[6]]
```

```{r 23_col_length_vs_nrows, exercise=TRUE}

```

```{r 23_col_length_vs_nrows-solution}
length(mpg_col)
length(col_6)
nrow(mtcars)
```

```{r 23_col_length_vs_nrows-code-check}
grade_code()
```

If each column in the `data.frame` is a `vector`, then the length of the vectors is the number of rows in the `data.frame`. Make sense?

### Summary statistics on columns

You can calculate summary statistic on column to characterize the data in `mtcars`.

Try these:

- display the mean displacement for all of the cars in `mtcars`
- display the range in fuel economy for the cars in `mtcars` (hint: `?range`)
- display a vector different number of cylinders that the cars might have (hint: `?unique`) 

```{r 24_mean_disp, exercise=TRUE}

```

```{r 24_mean_disp-solution}
mean(mtcars$disp)
range(mtcars$mpg)
unique(mtcars$cyl)
```

```{r 24_mean_disp-code-check}
grade_code()
```

## Indexing `data.frames`

`Data.frames` are lists.  Therefore, they can be indexed with `[`.  In other words, you can extract a subset of a `data.frame` using the `[` operator.  Because they are lists, they can be indexed with a single vector, just as you have practiced.  Give it a try:

- use `[` to subset mtcars.  Include columns 1, 4, and 6.  Assign the result to `three_column_mtcars`.
- display the first four rows of `three_column_mtcars`

```{r 25_index_mtcars_as_list_1, exercise=TRUE}

```

```{r 25_index_mtcars_as_list_1-solution}
three_column_mtcars <- mtcars[c(1,4,6)]
head(three_column_mtcars, 4)
```

```{r 25_index_mtcars_as_list_1-code-check}
grade_code()
```

Try again:

- Use `[` to subset mtcars.  Include columns "mpg", "hp", and "wt".  Assign the result to `three_column_mtcars`.
- display the first *five* rows of `three_column_mtcars`


```{r 26_index_mtcars_as_list_2, exercise=TRUE}

```

```{r 26_index_mtcars_as_list_2-solution}
three_column_mtcars <- mtcars[c("mpg", "hp", "wt")]
head(three_column_mtcars, 5)
```

```{r 26_index_mtcars_as_list_2-code-check}
grade_code()
```

One more time:

- Use `[` to subset mtcars.  Include columns the columns corresponding to TRUE in the following list: TRUE, FALSE, FALSE, TRUE, FALSE, TRUE, FALSE, FALSE, FALSE, FALSE, FALSE.  Assign the result to `three_column_mtcars`.
- display the first *three* rows of `three_column_mtcars`


```{r 27_index_mtcars_as_list_3, exercise=TRUE}

```

```{r 27_index_mtcars_as_list_3-solution}
three_column_mtcars <- 
  mtcars[c(TRUE, FALSE, FALSE, TRUE, FALSE, TRUE, FALSE, FALSE, FALSE, FALSE, FALSE)]
head(three_column_mtcars, 3)
```

```{r 27_index_mtcars_as_list_3-code-check}
grade_code()
```

Remember, however, that `data.frames` have two dimensions, represented by rows and columns.  So, they can also be indexed in two dimension by putting TWO indexing vectors inside `[ ]` -- the first index subsets the rows and the second subsets the columns.  

For instance, we can take a subset of the only those cars that have 4 cylinder engines and keep only the first three columns.

```{r, echo=TRUE}
# remember, there are 32 rows in mtcars so `has_four_cyl` will be
# a logical vector of length 32!  If you don't remember what `==`
# does, in the console, use ?`==` (include the backticks.  The
# backtick key is in the upper left of the keyboard, below
# <escape>)
has_four_cyl <- mtcars$cyl == 4
has_four_cyl
mtcars[has_four_cyl, 1:3]
```

Now you try!

- create a logical vector called `has_poor_fuel_economy` that contains TRUE if a car's fuel economy is less than 15 miles per gallon
- index the `mtcars` data.frame so the result contains the fuel economy, number of cylinders, horsepower, weight, and number of forward gears for cars with poor fuel economy.

```{r 28_index_mtcars_2D_1, exercise=TRUE}

```

```{r 28_index_mtcars_2D_1-solution}
has_poor_fuel_economy <- mtcars$mpg < 15
mtcars[has_poor_fuel_economy, c("mpg", "cyl", "hp", "wt", "gear")]
```

```{r 28_index_mtcars_2D_1-code-check}
grade_code()
```

If you want to include *all* of the columns but still filter the rows, you can leave out the index for the rows, but *be sure to include the comma* to R knows to use two dimensional indexing!  For instance, the following displays all columns for the car models with four cylinder engines.

```{r, echo=TRUE}
mtcars[has_four_cyl, ]
```

You try:

- display all of the columns for cars that have poor fuel economy. 

```{r 29_index_mtcars_2D_2-setup}
has_poor_fuel_economy <- mtcars$mpg < 15
```

```{r 29_index_mtcars_2D_2, exercise=TRUE}

```

```{r 29_index_mtcars_2D_2-solution}
mtcars[has_poor_fuel_economy, ]
```

```{r 29_index_mtcars_2D_2-code-check}
grade_code()
```

You can also leave out the row index to include all the rows.  For instance:

```{r, echo=TRUE}
mtcars[, c("mpg", "cyl", "hp", "wt", "gear")]
```

Yields all rows, but only the specified columns. Of course this is the same as indexing with a single vector:

```{r}
identical(
  # return spcific columns but keep all rows
  mtcars[ , c("mpg", "cyl", "hp", "wt", "gear")],
  # return entire columns
  mtcars[c("mpg", "cyl", "hp", "wt", "gear")]
)
```

### Putting it tegether

"Displacement" is the sum of volume of all cylinders in a car.  Therefore, we'd expect more displacement for cars with more cylinders in their engines.

Try this... You can index the rows of `mtcars` and then use the `$` function to extract a column.  The column will contain date from only those rows selected by the index.  By indexing `mtcars` and then extracting a column... 

- find the mean displacement for 4 cylinder engines
- find the mean displacement for 6 cylinder engines
- find the mean displacement for 8 cylinder engines

```{r 30_index_then_extract, exercise=TRUE}

```

```{r 30_index_then_extract-solution}
mean(mtcars[mtcars$cyl == 4, ]$disp)
mean(mtcars[mtcars$cyl == 6, ]$disp)
mean(mtcars[mtcars$cyl == 8, ]$disp)
```

```{r 30_index_then_extract-code-check}
grade_code()
```

### A note about the `[` operator (indexing)

By default, when you index a data.frame by row and column, but ask for only one column, you get a `vector` in return.  Notice the result of the following code is not a data.frame:

```{r, echo=TRUE}
mtcars[mtcars$cyl == 4, "disp"]
```

This is unexpected. 

The rule is that `[` should always return the same type of object that is being indexed.  Therefore, if you index a `data.frame`, you should get a `data.frame`

There is another type of `data.frame` called a `tibble` that follows the rules.  You will always get a `tibble` when you index a `tibble` -- always.  We'll learn more about `tibbles` later on, but for now, let's convert mtcars to a tibble and try two different ways of extracting a column from a subset of rows in a `data.frame.`


```{r, echo=TRUE}
# make a copy of mtcars as a tibble.
mtcars_tibble <- tibble::as_tibble(mtcars)

# compare results when we select several rows and *one* 
# column using indexing...
identical(
  # index one column from mtcars yields a vector
  mtcars[mtcars$ccyl == 4, "mpg"],
  # index one column from mtcars_tibble yields a 
  # one-column data.frame
  mtcars_tibble[mtcars_tibble$cyl == 4, "mpg"]
)

# select some rows across all columns, then extract a column...
identical(
  # from mtcars yields a vector
  mtcars[mtcars$cyl == 4, ]$mpg,
  # from mtcars_tibble also yields a vector
  mtcars_tibble[mtcars_tibble$cyl == 4, ]$mpg
)
```

Your code with be more consistent and therefore less error prone if you use the second method -- use *indexing* to select the rowse you want and then *extract* the column of interest!

### Nice work!

Turn in your homework using `ENSC311::submit_ENSC311()`.  When you do so, look carefully at the list of options to turn in.  They will be in alphabetical not chronological order... 

Be sure to choose the number corresponding to the homework for Lecture 8! 
