---
title: "test tutorialize"
output: learnr::tutorial
runtime: shiny_prerendered
---
<!-- File created by tutorialize.  Do not edit by hand. -->

```{r setup, include=FALSE}
library(learnr)
knitr::opts_chunk$set(echo = FALSE)
gradethis::gradethis_setup()
```


## Topic 1

### Exercise 

*Here's a simple exercise with an empty code chunk provided for entering the answer.*

Write the R code required to add two plus two:

```{r two-plus-two-setup} 
x = 2
```

```{r two-plus-two, exercise=TRUE}

```

```{r two-plus-two-solution}
y = 2
answer <- x + y
answer
```

```{r two-plus-two-code-check}
grade_code()
```

### Exercise with Code

*Here's an exercise with some prepopulated code as well as `exercise.lines = 5` to provide a bit more initial room to work.*

Now write a function that adds any two numbers and then call it with x and y:


```{r add-function-setup}
x = 2
y = 2
answer <- x + y
z <- 3
```

```{r add-function, exercise=TRUE, exercise.lines=5}
add <- function(____, ____, ____) {

}
```

```{r add-function-solution}
add <- function(a, b, c) { 
  a + b + c
}
add(x, y, z)
```

```{r add-function-code-check}
grade_code()
```

Now call `add()` with the literals for 5 and 10.

```{r add-function_test-setup}
x = 2
y = 2
answer <- x + y
z <- 3
add <- function(a, b, c) {
    a + b + c
}
```

```{r add-function_test, exercise=TRUE, exercise.lines=5}

```

```{r add-function_test-solution}
add(5, 10, 15)
```

```{r add-function_test-code-check}
grade_code()
```


## Topic 2

### Exercise with Hint

*Here's an exercise where the chunk is pre-evaluated via the `exercise.eval` option (so the user can see the default output we'd like them to customize). We also add a "hint" to the correct solution via the chunk immediate below labeled `print-limit-hint`.*

Modify the following code to limit the number of rows printed to 5:

```{r print-limit, exercise=TRUE}

```

```{r print-limit-solution}
mtcars
```

```{r print-limit-code-check}
grade_code()
```



### Quiz



*You can include any number of single or multiple choice questions as a quiz. Use the `question` function to  define a question and the `quiz` function for grouping multiple questions together.*

```{r}
x=4
```

Some questions to verify that you understand the purposes of various base and recommended R packages:

```{r quiz}
quiz(
  question("Which package contains functions for installing other R packages?",
    answer("base"),
    answer("tools"),
    answer("utils", correct = TRUE),
    answer("codetools")
  ),
  question("Which of the R packages listed below are used to create plots?",
    answer("lattice", correct = TRUE),
    answer("tools"),
    answer("stats"),
    answer("grid", correct = TRUE)
  )
)
```
